```scala
import $file.hw4stdlib
import hw4stdlib._
```

Output:
```
<console>:90: error: not found: value hw4stdlib
       import hw4stdlib._
              ^
```

```scala

// Nats

sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one = Succ(Zero)
val two = Succ(one)
val three = Succ(two)
val four = Succ(three)
val five = Succ(four)
val six = Succ(five)
val seven = Succ(six)
val eight = Succ(seven)
val nine = Succ(eight)
val ten = Succ(nine)

def nat_plus(x : Nat, y : Nat) : Nat = x match {
    case Zero    => y
    case Succ(x) => Succ(nat_plus(x, y))
}

def nat_to_int(x : Nat) : Int = x match {
    case Zero => 0
    case Succ( x ) => 1 + nat_to_int(x)
}

def nat_to_str(x : Nat) = nat_to_int(x).toString()

def print_nat(x : Nat) = println(nat_to_str(x))

def nat_mult(x : Nat, y : Nat) : Nat = x match {
    case Zero    => Zero
    case Succ(x) => nat_plus(nat_mult(x,y), y)
}


def nat_minus(x : Nat, y : Nat) : Nat = (x, y) match {
    case (Zero, _)          => Zero
    case (x, Zero)          => x
    case (Succ(x), Succ(y)) => nat_minus(x, y)
}

def nat_pow(x : Nat, y : Nat) : Nat = x match {
    case Zero       => one
    case Succ(Zero) => y 
    case Succ(x)    => nat_mult(y, nat_pow(x, y))
}

def nat_lte(x : Nat, y : Nat) : Bool = (x, y) match {
    case (Zero, y)          => True
    case (x, Zero)          => False
    case (Succ(x), Succ(y)) => nat_lte(x, y)
}

// Booleans

sealed trait Bool
case object True extends Bool
case object False extends Bool

def t = True
def f = False

def id(x : Bool) : Bool = x

def not(x : Bool) : Bool = x match {
    case True => False
    case False => True
}

def and(x : Bool, y : Bool) : Bool = (x,y) match {
    case (True, True) => True
    case _ => False
}

def or(x : Bool, y : Bool) : Bool = (x, y) match {
    case (False, False) => False
    case _              => True
}

def xor(x : Bool, y : Bool) : Bool = (x, y) match{
    case (True, False) => True
    case (False, True) => True
    case _ => False
}

def nand(x : Bool, y : Bool) : Bool = not(and(x,y))

def lte(x: Nat, y: Nat): Bool = (x, y) match {
    case (Zero, _) => True
    case (_, Zero) => False
    case (Succ(p1), Succ(p2)) => lte(p1, p2)
}

// Lists

sealed trait List[+A]
case object Empty extends List[Nothing]
case class Cons[A](x : A, xs : List[A]) extends List[A]

def length[A](xs : List[A]) : Nat = xs match {
    case Empty => Zero
    case Cons(_, xs) => Succ(length(xs))
}

def append[A](xs : List[A], ys : List[A]) : List[A] = xs match {
    case Empty => ys
    case Cons(x, xs) => Cons(x, append(xs, ys))
}

def fold[A, B](f : ((A, B) => B), acc : B, xs : List[A]) : B = xs match {
    case Empty       => acc
    case Cons(x, xs) => fold(f, f(x, acc), xs)
}

def map[A,B](f : (A => B), xs : List[A]) : List[B] = xs match {
    case Empty => Empty
    case Cons(x, xs) => Cons(f(x), map(f,xs))
}

def filter[A](p : (A => Bool), xs : List[A]) : List[A] = xs match {
    case Empty       => Empty
    case Cons(x, xs) => p(x) match {
        case True  => Cons(x, filter(p, xs))
        case False => filter(p, xs)
    }
}

def reverse[A](xs : List[A]) : List[A] =
    fold(Cons(_: A, _: List[A]), Empty, xs)

// Maybe
sealed trait Maybe[+A]
case object None extends Maybe[Nothing]
case class Just[A](x : A) extends Maybe[A]

def map_maybe[A,B](f : (A => B), mx : Maybe[A]) : Maybe[B] = mx match {
    case None => None
    case Just(x) => Just(f(x))
}

// Products

sealed trait Pair[+A, +B]
case class MkPair[A, B](fst : A, snd : B) extends Pair[A, B]

// Sums

sealed trait Either[+A, +B]
case class Left[A, B](left : A) extends Either[A, B]
case class Right[A, B](right : B) extends Either[A, B]

// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

```

Output:
```
defined trait Nat
defined object Zero
defined class Succ
one: Succ = Succ(Zero)
two: Succ = Succ(Succ(Zero))
three: Succ = Succ(Succ(Succ(Zero)))
four: Succ = Succ(Succ(Succ(Succ(Zero))))
five: Succ = Succ(Succ(Succ(Succ(Succ(Zero)))))
six: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))
seven: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))
eight: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))
nine: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))))
ten: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))))
nat_plus: (x: Nat, y: Nat)Nat
nat_to_int: (x: Nat)Int
nat_to_str: (x: Nat)String
print_nat: (x: Nat)Unit
nat_mult: (x: Nat, y: Nat)Nat
nat_minus: (x: Nat, y: Nat)Nat
nat_pow: (x: Nat, y: Nat)Nat
nat_lte: (x: Nat, y: Nat)Bool
define...
```

# Homework 4

This is a longer assignment due to the exam.
Due 10/3 at 11:59pm

## Submission Instructions
Upload only this .ipynb file to Canvas. Do not add anything to stdlib since you can't submit it.

In this homework we will define integers and their operations, then build an interpreter.

See [This link](https://www.notion.so/Guidelines-for-Programming-Homework-dbd25efa7bb24915ae6bcb06827fc5b6) for what is and isn't allowed in your code. 

## Problem 1 (20 pts)

For the arithmetic interpreter we will write later on in the homework we will need arithmetic operations for the integers. These will obviously correspond very closely to the operations defined on the Natural Numbers. But now we will need to account for negative numbers in a proper way. 

Start by defining the Algebraic Datatype(sealed trait with case classes) for the Integers. The encoded type should be:

$$
\begin{align}
    \mathbb{Z} ::=&\ \text{Positive}\ \mathbb{N}\\
    \mid&\ \text{Negative}\ \mathbb{N}
\end{align}
$$

We will use this encoding because it will allow us to use some of our definitions for the Natural Numbers to define the operations on the Integers. This is the core strategy of functional programming: using smaller functions to build up larger ones. This is true even across types. We want to always minimize the ammount of code we have be reusing what we have already constructed, whenever possible. 

### 1a

Define the ADT for Integers below: 

<!-- grade_id: cell-9ee3e730ef837143, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
sealed trait Integer
case class Positive(x : Nat) extends Integer
case class Negative(x : Nat) extends Integer
```

Output:
```
defined trait Integer
defined class Positive
defined class Negative
```

<!-- grade_id: cell-cdc9d6798e91c35c, points: 3, grade: true, solution: false -->
```scala
def ione : Integer = Positive(Succ(Zero))
def ineg_one : Integer = Negative(Succ(Zero))

def int_to_str(x : Integer) : String = x match {
    case Positive(x) => nat_to_str(x)
    case Negative(x) => "-" + nat_to_str(x)
}

def print_integer(x : Integer) = println(int_to_str(x))

print_integer(ione)
print_integer(ineg_one)
passed(3)
```

Output:
```
1
-1

*** Tests Passed (3 points) ***
null
```

### Absolute Value and Negation

Below are defined the absolute value function and a negation function

```scala
def abs(x : Integer) : Nat = x match {
    case Positive(x) => x
    case Negative(x) => x
}

def negate(x : Integer) : Integer = x match {
    case Positive(x) => Negative(x)
    case Negative(x) => Positive(x)
}
```

Output:
```
abs: (x: Integer)Nat
negate: (x: Integer)Integer
```

### 1b: Addition

Define plus and minus for Integers. Don't use the versions we created for $\mathbb{N}$ as it did some funky things to get minus to behave correctly. Try and create these from scratch instead. It is recommended to use the lte(less than or equals) function we defined last week along with abs to make the job easier. 

Note that we have renamed the operations for natural numbers so that they are of the form:

``` nat_<operation name> ```

For instance, plus has been renamed to ``` nat_plus ```. This was done so we don't have namespace conflicts

<!-- grade_id: cell-3f813578972946cd, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def plus(x : Integer, y : Integer) : Integer = (x, y) match {
    case (Positive(a), Positive(b)) => Positive(nat_plus(abs(x), abs(y)))
    case (Negative(a), Negative(b)) => Negative(nat_plus(abs(x), abs(y)))
    case (_, Negative(b)) => nat_lte(abs(x), abs(y)) match {
        case True => Negative(nat_minus(abs(y), abs(x)))
        case False => Positive(nat_minus(abs(x), abs(y)))
    }
    case (_, Positive(b)) => nat_lte(abs(x), abs(y)) match {
        case True => Positive(nat_minus(abs(y), abs(x)))
        case False => Negative(nat_minus(abs(x), abs(y)))
    }
}

```

Output:
```
plus: (x: Integer, y: Integer)Integer
```

<!-- grade_id: cell-816461e09874b44b, points: 5, grade: true, solution: false -->
```scala
assert(plus(Positive(three), Negative(six)) == Negative(three))
assert(plus(Positive(three), Negative(two)) == Positive(one))
assert(plus(Positive(five), Positive(three)) == Positive(eight))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### 1c: Subtraction
Implement subtraction below
Hint: Subtraction is very easy if you use ```plus``` and a ```negate```

<!-- grade_id: cell-cdaed49e10461f7f, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def minus(x : Integer, y : Integer) : Integer = (x, y) match {
    case (Positive(x), Negative(y)) => plus(Positive(x), negate(Negative(y)))
    case (Negative(x), Positive(y)) => plus(Negative(x), negate(Positive(y)))
    case (Positive(x), Positive(y)) => plus(Positive(x), negate(Positive(y)))
    case (Negative(x), Negative(y)) => plus(Negative(x), negate(Negative(y)))
}
```

Output:
```
minus: (x: Integer, y: Integer)Integer
```

<!-- grade_id: cell-2f31b73a07aa01f8, points: 3, grade: true, solution: false -->
```scala
assert(minus(Positive(three), Negative(six)) == Positive(nine))
assert(minus(Positive(three), Negative(two)) == Positive(five))
assert(minus(Positive(five), Positive(six)) == Negative(one))
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

### 1d: Multiplication

Write multiplication for Integers. You should be able to use ```nat_mult``` to greatly simplify this

<!-- grade_id: cell-1fb42ea8405c27ad, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def mult(x : Integer, y : Integer) : Integer = (x, y) match {
    case (Positive(a), Positive(b)) => Positive(nat_mult(abs(x), abs(y)))
    case (Negative(a), Negative(b)) => Positive(nat_mult(abs(x), abs(y)))
    case _ => Negative(nat_mult(abs(x), abs(y)))
}
```

Output:
```
mult: (x: Integer, y: Integer)Integer
```

<!-- grade_id: cell-ab71ea516b48a773, points: 5, grade: true, solution: false -->
```scala
assert(mult(Positive(three), Negative(two)) == Negative(six))
assert(mult(Positive(two), Positive(one)) == Positive(two))
assert(mult(Negative(three), Negative(three)) == Positive(nine))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### 1e: Exponentiation

Recall that for ``pow`` we will restrict ourselves to only positive powers. Use the definition of `mult` from above so that your polarity(Positive/Negative) is correct. Recall the cases for $-x^n$ for even vs odd $n$. A hint for your base case: `pow(x, 0) = Positive( 1 )`

<!-- grade_id: cell-8bf97e036b7c5e3c, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def pow(x : Integer, p : Nat) : Integer = (x, p) match {
    case (x, Zero) => Positive(one)
    case (_, Succ(y)) => mult(x, pow(x, y))
}
```

Output:
```
pow: (x: Integer, p: Nat)Integer
```

<!-- grade_id: cell-8bfc201c0fa1c3ab, points: 4, grade: true, solution: false -->
```scala
assert(pow(Negative(two), two) == Positive(four))
assert(pow(Positive(three), one) == Positive(three))
assert(pow(Negative(two), three) == Negative(eight))
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

# The Arithmetic Language: Our First Interpreter

Now we are ready to define our first interpreter. We will define the Arithmetic language syntax below as a sealed trait. It will be your job to correctly construct the interpreter for it based on the inference rules we covered in class. Recall that each rule corresponds to a case in the `eval` function.

```scala
sealed trait Expr 
case class Num(x : Integer) extends Expr 
case class Plus(x : Expr, y : Expr) extends Expr
case class Minus(x : Expr, y : Expr) extends Expr
case class Mult(x : Expr, y : Expr) extends Expr
case class Pow(x : Expr, y : Nat) extends Expr
```

Output:
```
defined trait Expr
defined class Num
defined class Plus
defined class Minus
defined class Mult
defined class Pow
```

## Problem 2 (10 points)

Now that we have defined the syntax for Arithmetic expressions. Go ahead and define the interpreter. We have given the signature for the function. (Bonus points if you define and use the helper function `eval-bin` that we discussed in class.

Recall that the type of this interpreter should be $eval : \text{Expr} \rightarrow \mathbb{Z}$

<!-- grade_id: cell-cd1a5c1ec38ae400, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def eval_bin(f : (Integer, Integer) => Integer, e1 : Expr, e2 : Expr) = (e1, e2) match {
    case (_, _) => f(eval(e1), eval(e2))
}

def eval(x : Expr) : Integer = x match {
    case Num(n) => n
    case Plus(n, m) => eval_bin(plus, n, m)
    case Minus(n, m) => eval_bin(minus, n, m)
    case Mult(n, m) => eval_bin(mult, n, m)
    // Pow does not take (Int, Int) so eval_bin won't work
    case Pow(n, m) => pow(eval(n), m)
}
```

Output:
```
eval_bin: (f: (Integer, Integer) => Integer, e1: Expr, e2: Expr)Integer
eval: (x: Expr)Integer
```

<!-- grade_id: cell-14ceee2b601c4ef2, points: 4, grade: true, solution: false -->
```scala
val x: Expr = Num(Positive(six))
assert(eval(x) == Positive(six))
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
null
```

<!-- grade_id: cell-5ab79956b85cc902, points: 3, grade: true, solution: false -->
```scala
val x2: Expr = Plus(Num(Positive(two)), Num(Positive(two)))
assert(eval(x2) == Positive(four))
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
null
```

<!-- grade_id: cell-6b5e757e2dd296d2, points: 3, grade: true, solution: false -->
```scala
val x3: Expr = Mult(Plus(Num(Positive(two)), Num(Positive(two))), Num(Negative(two)))
assert(eval(x3) == Negative(eight))
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
null
```

## Problem 3 (10 points)
Implement equality for $\mathbb{B}$, $\mathbb{N}$, $\mathbb{Z}$, and $\text{List}\ a$

Most should have the form:
$$
eq : A \rightarrow A \rightarrow \mathbb{B}
$$

Where you will want to fill each $A$ with the type of equality you are defining

### 3a: $\mathbb{B}$
Implement `bool_eq`

<!-- grade_id: cell-2798ee2b3f2a50f6, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def bool_eq(x : Bool, y : Bool) : Bool = (x,y) match {
    case (True, True) => True
    case (False, False) => True
    case (_, _) => False
}
```

Output:
```
bool_eq: (x: Bool, y: Bool)Bool
```

<!-- grade_id: cell-722112b63bcac170, points: 2, grade: true, solution: false -->
```scala
assert(bool_eq(True, True) == True)
assert(bool_eq(False, True) == False)
assert(bool_eq(True, False) == False)
assert(bool_eq(False, False) == True)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
```

### 3b: $\mathbb{N}$
Implement `nat_eq`

<!-- grade_id: cell-22169f8784478436, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def nat_eq(x: Nat, y : Nat) : Bool = (x, y) match {
    case (Zero, Zero) => True
    case (_, Zero) => False
    case (Zero, _) => False
    case (Succ(a), Succ(b)) => nat_eq(a, b)
}
```

Output:
```
nat_eq: (x: Nat, y: Nat)Bool
```

<!-- grade_id: cell-0ce38fa2255b8c9d, points: 2, grade: true, solution: false -->
```scala
assert(nat_eq(ten, ten) == True)
assert(nat_eq(ten, Zero) == False)
assert(nat_eq(five, six) == False)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

### 3c: $\mathbb{Z}$
Implement `int_eq`

<!-- grade_id: cell-eab55be5227ceeaf, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def int_eq(x : Integer, y : Integer) : Bool = (x, y) match {
    case (Positive(Zero), Negative(Zero)) => True
    case (Negative(Zero), Positive(Zero)) => True
    case (Positive(a), Negative(b)) => False
    case (Negative(a), Positive(b)) => False
    case (Positive(a), Positive(b)) => nat_eq(a, b)
    case (Negative(a), Negative(b)) => nat_eq(a, b)
}
```

Output:
```
int_eq: (x: Integer, y: Integer)Bool
```

<!-- grade_id: cell-f846b69e86a7d77f, points: 2, grade: true, solution: false -->
```scala
assert(int_eq(Positive(nine), Positive(nine)) == True)
assert(int_eq(Negative(eight), Negative(eight)) == True)
assert(int_eq(Positive(nine), Negative(nine)) == False)
assert(int_eq(Positive(nine), Positive(Zero)) == False)
assert(int_eq(Positive(five), Positive(six)) == False)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

### 3d: $\text{List}\ a$
Implement `list_eq`. Since lists are polymorphic, your function needs to take a third parameter which should be the `eq` function for the given $a$:
$$
eq : \text{List}\ A \rightarrow \text{List}\ A \rightarrow (A \rightarrow A \rightarrow \mathbb{B}) \rightarrow \mathbb{B}
$$

<!-- grade_id: cell-7459289e4fe0f331, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def list_eq[A](x : List[A], y : List[A], eq : (A, A) => Bool) : Bool = (x, y) match {
    case (Empty, Empty) => True
    case (Empty, _) => False
    case (_, Empty) => False
    case (Cons(a, as), Cons(b, bs)) => eq(a, b) match {
        case True => list_eq(as, bs, eq)
        case False => False
    }
}
```

Output:
```
list_eq: [A](x: List[A], y: List[A], eq: (A, A) => Bool)Bool
```

<!-- grade_id: cell-39a05cd0083a25bd, points: 4, grade: true, solution: false -->
```scala
assert(list_eq(Empty, Empty, nat_eq) == True)
assert(list_eq(Empty, Empty, bool_eq) == True)
assert(list_eq(Cons(True, Empty), Cons(True, Empty), bool_eq) == True)
assert(list_eq(Cons(True, Empty), Empty, bool_eq) == False)
assert(list_eq(Cons(True, Cons(False, Empty)), Cons(True, Cons(True, Empty)), bool_eq) == False)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

```scala

```